## Tips
- To avoid unpleasant surprises, I suggest you _run all cells in their order of appearance_ (__Cell__ $\rightarrow$ __Run All__).


- If the changes you've made to your solution don't seem to be showing up, try running __Kernel__ $\rightarrow$ __Restart & Run All__ from the menu.


- Before submitting your assignment, make sure everything runs as expected. First, restart the kernel (from the menu, select __Kernel__ $\rightarrow$ __Restart__) and then **run all cells** (from the menu, select __Cell__ $\rightarrow$ __Run All__).

## Reminder

- Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name, UA email, and collaborators below:



Several of the cells in this notebook are **read only** to ensure instructions aren't unintentionally altered.  

If you can't edit the cell, it is probably intentional.

```python
NAME = "Wenmo Sun"
# University of Arizona email address
EMAIL = "wmsun@email.arizona.edu"
# Names of any collaborators.  Write N/A if none.
COLLABORATORS = "N/A"
```

## Scratchpad

You are welcome to create new cells (see the __Cell__ menu) to experiment and debug your solution.

```python
%load_ext autoreload
%autoreload 2
```

# Mini Python tutorial

This course uses Python 3.8.

Below is a very basic (and incomplete) overview of the Python language... 

For those completely new to Python, [this section of the official documentation may be useful](https://docs.python.org/3.8/library/stdtypes.html#common-sequence-operations).

```python
# This is a comment.  
# Any line starting with # will be interpreted as a comment

# this is a string assigned to a variable
greeting = "hello"

# If enclosed in triple quotes, strings can also be multiline:

"""
I'm a multiline
string.
"""

# let's use a for loop to print it letter by letter
for letter in greeting:
    print(letter)
    
# Did you notice the indentation there?  Whitespace matters in Python!

# here's a list of integers

numbers = [1, 2, 3, 4]

# let's add one to each number using a list comprehension
# and assign the result to a variable called res
# list comprehensions are used widely in Python (they're very Pythonic!)

res = [num + 1 for num in numbers]

# let's confirm that it worked
print(res)

# now let's try spicing things up using a conditional to filter out all values greater than or equal to 3...
print([num for num in res if not num >= 3])

# Python 3.7 introduced "f-strings" as a convenient way of formatting strings using templates
# For example ...
name = "Josuke"

print(f"{greeting}, {name}!")

# f-strings are f-ing convenient!


# let's look at defining functions in Python..

def greet(name):
    print(f"Howdy, {name}!")

# here's how we call it...

greet("partner")

# let's add a description of the function...

def greet(name):
    """
    Prints a greeting given some name.
    
    :param name: the name to be addressed in the greeting
    :type name: str
    
    """
    print(f"Howdy, {name}!")
    
# I encourage you to use docstrings!

# Python introduced support for optional type hints in v3.5.
# You can read more aobut this feature here: https://docs.python.org/3.8/library/typing.html
# let's give it a try...
def add_six(num: int) -> int:
    return num + 6

# this should print 13
print(add_six(7))

# Python also has "anonymous functions" (also known as "lambda" functions)
# take a look at the following code:

greet_alt = lambda name: print(f"Hi, {name}!")

greet_alt("Fred")

# lambda functions are often passed to other functions
# For example, they can be used to specify how a sequence should be sorted
# let's sort a list of pairs by their second element
pairs = [("bounce", 32), ("bighorn", 12), ("radical", 4), ("analysis", 7)]
# -1 is last thing in some sequence, -2 is the second to last thing in some seq, etc.
print(sorted(pairs, key=lambda pair: pair[-1]))

# we can sort it by the first element instead
# NOTE: python indexing is zero-based
print(sorted(pairs, key=lambda pair: pair[0]))

# You can learn more about other core data types and their methods here: 
# https://docs.python.org/3.8/library/stdtypes.html

# Because of its extensive standard library, Python is often described as coming with "batteries included".  
# Take a look at these "batteries": https://docs.python.org/3.8/library/

# You now know enough to complete this homework assignment (or at least where to look)
```

Output:
```
h
e
l
l
o
[2, 3, 4, 5]
[2]
hello, Josuke!
Howdy, partner!
13
Hi, Fred!
[('radical', 4), ('analysis', 7), ('bighorn', 12), ('bounce', 32)]
[('analysis', 7), ('bighorn', 12), ('bounce', 32), ('radical', 4)]
```

```python

```

# Bonus

These tests cover the bonus exercise and are meant for students looking for challenges and/or extra credit.

```python
import itertools
import numpy as np

from ipynb.fs.full.assignment import Classifier, read_ptbtagged
```

Output:
```
File "<frozen importlib._bootstrap_external>", line 779, in exec_module

  File "/usr/local/lib/python3.8/site-packages/ipynb/fs/full/__init__.py", line 43, in get_code
    return self.source_to_code(code_from_ipynb(nb), self.path)

  File "/app/assignment.ipynb", line 37
    "If you can't edit the cell, it is probably intentional."
    ^
SyntaxError: invalid syntax
```

```python
np.random.seed(42)
```

```python
# part-of-speech tags from the Penn Treebank
PTB_TAGS = {
    "#", "$", "''", "``", ",", "-LRB-", "-RRB-", ".", ":", "CC", "CD", "DT",
    "EX", "FW", "IN", "JJ", "JJR", "JJS", "LS", "MD", "NN", "NNP", "NNPS",
    "NNS", "PDT", "POS", "PRP", "PRP$", "RB", "RBR", "RBS", "RP", "SYM", "TO",
    "UH", "VB", "VBD", "VBG", "VBN", "VBP", "VBZ", "WDT", "WP", "WP$", "WRB",
}
```

## BONUS: Viterbi decoding (+5 pts)

Try your hand at implementing Viterbi decoding to find the optimal tag sequence based on MEMM predictions for a given observation sequence.  

NOTE: You'll need to implement `Classifier.predict_viterbi` in `assignment.ipynb`.

<!-- grade_id: bonus-viterbi-decoding, points: 5, grade: true, solution: false -->
```python
def test_predict_viterbi():
    clf        = Classifier()
    ptb_train  = read_ptbtagged("data/PTBSmall/train.tagged")
    ptb_train  = itertools.islice(ptb_train, 2)  # just the 1st 2 sentences
    clf.train(ptb_train)

    # POS tags in first 2 sentences
    possible_tags = {"NNP", ",", "CD", "NNS", "JJ", "MD", "VB", "DT", "NN",
                     "IN", ".", "VBZ", "VBG"}
    n_tags = len(possible_tags)

    # sample sentence to be fed to classifier
    tokens = "Vinken is a director .".split()
    trans_probs, viterbi_lattice, pos_tags = clf.predict_viterbi(tokens)

    # check that the transition probabilities are the right shape
    # second axis is +1 since the last entry is transitions starting at <s>
    assert trans_probs.shape == (len(tokens), n_tags + 1, n_tags)

    # check that probability distribution from <s> to first word sums to 1
    s_index = n_tags
    prob_dist_sums = np.sum(np.exp(trans_probs), axis=-1)
    np.testing.assert_almost_equal(prob_dist_sums[0, s_index], 1)

    # check that probability distributions of all non-<s> pairs sum to 1
    for i in range(1, len(tokens)):
        for j in range(0, n_tags):
            np.testing.assert_almost_equal(prob_dist_sums[i, j], 1)

    # check that the lattice is the right shape
    assert viterbi_lattice.shape == (len(tokens), n_tags)

    # check that the numbers are not all the same in the lattice
    assert np.std(viterbi_lattice) > 0
    assert np.all(np.std(viterbi_lattice, axis=0) > 0)
    assert np.all(np.std(viterbi_lattice, axis=1) > 0)

    # check that the probabilities from <s> are on the first token
    np.testing.assert_almost_equal(viterbi_lattice[0], trans_probs[0, s_index])

    # check some probability calculations in the lattice
    np.testing.assert_almost_equal(viterbi_lattice[1, 2], max([
        viterbi_lattice[0, k] + trans_probs[1, k, 2] for k in range(n_tags)]))
    np.testing.assert_almost_equal(viterbi_lattice[3, 1], max([
        viterbi_lattice[2, k] + trans_probs[3, k, 1] for k in range(n_tags)]))
    np.testing.assert_almost_equal(viterbi_lattice[-1, 9], max([
        viterbi_lattice[-2, k] + trans_probs[-1, k, 9] for k in range(n_tags)]))

    # check that the POS tags are all valid tags
    assert all([pos_tag in PTB_TAGS for pos_tag in pos_tags])

    # check that the lattice's score for the predicted POS tag path matches
    # the score we would get from the transition probabilities
    pos_indexes = [clf.label_index(t) for t in pos_tags]
    np.testing.assert_almost_equal(
        viterbi_lattice[-1, pos_indexes[-1]],
        sum(trans_probs[i, pos_indexes[i - 1] if i else s_index, pos_index]
            for i, pos_index in enumerate(pos_indexes)))

    # check that the selected POS path has the highest score of all the possible
    # paths through the lattice
    np.testing.assert_almost_equal(
        viterbi_lattice[-1, pos_indexes[-1]],
        max(trans_probs[0, s_index, index1] +
            trans_probs[1, index1, index2] +
            trans_probs[2, index2, index3] +
            trans_probs[3, index3, index4] +
            trans_probs[4, index4, index5]
            for index1 in range(n_tags)
            for index2 in range(n_tags)
            for index3 in range(n_tags)
            for index4 in range(n_tags)
            for index5 in range(n_tags)))

test_predict_viterbi()
```

Output:
```
<ipython-input-12-67f25beff07d> in <module>
     73             for index5 in range(n_tags)))
     74 
---> 75 test_predict_viterbi()

<ipython-input-12-67f25beff07d> in test_predict_viterbi()
      1 def test_predict_viterbi():
----> 2     clf        = Classifier()
      3     ptb_train  = read_ptbtagged("data/PTBSmall/train.tagged")
      4     ptb_train  = itertools.islice(ptb_train, 2)  # just the 1st 2 sentences
      5     clf.train(ptb_train)

NameError: name 'Classifier' is not defined
```